# The McDonald's Diet

In this exercise, we will construct a variety of daily meal plans made up entirely of items from McDonald's. Let's to see what data I have.

## Problem 1

You first task is to read in McDonaldsmenu.csv data set and return a dataframe with  only the "Item" and "Calories" columns, as well as any column whose column name contains a "%".

**IMPORTANT: When you read in the data, you must use the relative file path, i.e. the path Data/McDonaldsmenu.csv.**


```python
import pandas as pd
from gurobipy import *
```

<!-- grade_id: cell-38cd49bb28c7b80f, points: 0, grade: false, solution: true -->
```python
def problem1():

    import numpy as np
    df_mcdonalds = pd.read_csv("Data/McDonaldsmenu.csv")
    nutrition = []
    for col_name in df_mcdonalds.columns:
        if "(% Daily Value)" in col_name:
            nutrition.append(col_name)
    nutrition.append("Calories")
    nutrition.append("Item")
    df_nutrition = df_mcdonalds[nutrition]
    return df_nutrition
```

```python
problem1()
```

Output:
```
Total Fat (% Daily Value)  Saturated Fat (% Daily Value)  \
0                           20                             25   
1                           12                             15   
2                           35                             42   
3                           43                             52   
4                           35                             42   
..                         ...                            ...   
255                         26                             44   
256                         35                             58   
257                         17                             29   
258                         50                             76   
259                         25                             38   

     Cholesterol (% Daily Value)  Sodium (% Daily Value)  \
0                             87                      31   
1                              8                      32   
2                             15            
```

<!-- grade_id: cell-460b84e6b86f0a5d, points: 5, grade: true, solution: false -->
```python
df = problem1()
assert df.shape == (260,12)
```

## Problem 2

 Let's create daily meal plan from these items that have to abide by the following constraints:

* Any particular item can only be ordered at most twice
* We have to satisfy the daily % values:
    - $\leq$ constraints for: Total Fat, Saturated Fat, Cholesterol, Sodium, Carbohydrates
    - $\geq$ constraint for the others
   
 

We will solve the problem with the following objective:

* Minimize calories

Return the number of calories in the optimal solution as well as a dictionary that has a key for each item ordered in the optimal solution, whose value is the number of times that item should be ordered.

**Hint: You should be calling your function from problem 1 to read in the data**

<!-- grade_id: cell-1e9cf24cb531b147, points: 0, grade: false, solution: true -->
```python
def problem2():
    
    df_nutrition = problem1()
    total_cals = None
    m = Model()
    order_dictionary = {}
    menu = {}
    length = 0 
    for index, row in df_nutrition.iterrows():
        item_name = df_nutrition.loc[index, 'Item']
        menu[item_name] = m.addVar(lb = 0, ub = 2, vtype = GRB.INTEGER, name = "x_%s" % item_name)
    m.setObjective(expr = quicksum(menu[item_name] * df_nutrition.loc[index, "Calories"]\
                                   for index, item_name in df_nutrition['Item'].iteritems()), sense = GRB.MINIMIZE)

    nutrients = ["Total Fat (% Daily Value)", "Saturated Fat (% Daily Value)", "Cholesterol (% Daily Value)", "Sodium (% Daily Value)", "Carbohydrates (% Daily Value)", \
                "Dietary Fiber (% Daily Value)", "Calcium (% Daily Value)", "Iron (% Daily Value)", "Vitamin C (% Daily Value)", "Vitamin A (% Daily Value)"]
    for nutrient in nutrients:
        if nutrient in df_nutrition.columns:
            if nutrient in ["Total Fat (% Daily Value)", "Saturated Fat (% Daily Value)", "Cholesterol (% Daily Value)", "Sodium (% Daily Value)", "Carbohydrates (% Daily Value)"]:
                m.addConstr(quicksum(menu[item_name] * df_nutrition.loc[index, nutrient] for index, item_name in df_nutrition['Item'].iteritems()) <= 100)
            else:
                m.addConstr(quicksum(menu[item_name] * df_nutrition.loc[index, nutrient] for index, item_name in df_nutrition['Item'].iteritems()) >= 100)
                            
    m.setParam( 'OutputFlag', False )
    m.optimize()
    
    total_cals = m.objVal
    
    for v in m.getVars():
        if v.X > 0:
            item_name = v.varName.split("_")[1]
            order_dictionary[item_name] = int(v.X)
    

    return total_cals, order_dictionary


```

```python
problem2()
```

Output:
```
Set parameter Username
Academic license - for non-commercial use only - expires 2024-09-12
(1640.0,
 {'Fruit & Maple Oatmeal without Brown Sugar': 1,
  'Hamburger': 2,
  'Premium Southwest Salad (without Chicken)': 2,
  'Side Salad': 2,
  'Chocolate Chip Cookie': 2,
  'Fat Free Chocolate Milk Jug': 2})
```

<!-- grade_id: cell-a924d475c1e34982, points: 10, grade: true, solution: false -->
```python
total_cals, order_dictionary = problem2()
assert  total_cals == 1640
```

<!-- grade_id: cell-5edb55b85470306b, points: 5, grade: true, solution: false -->
```python
total_cals, order_dictionary = problem2()
assert len(order_dictionary.keys()) == 6
assert order_dictionary['Premium Southwest Salad (without Chicken)'] == 2
```

```python

```